## Homework 5: Natural language processing

In this assignment, we'll practice reading PDF files, topic modeling, and sentiment analysis.

We'll use [this dataset of emailed comments](https://minneapolis2040.com/media/1411/minneapolis-2040-emailed-responses_062218.pdf) on the Minneapolis 2040 plan. It's in your repostitory as `minneapolis-2040-emailed-responses_062218.pdf`. If you want some background, [check out the city's website](https://minneapolis2040.com/received-public-comments/).

Please help us grade by observing the following:
 
* Do not rename this notebook (that messes up the autograder)
* Do not include large sections of output (that makes it hard to find your code). For example, use `df.head()` to show the first few rows, rather than printing an entire dataframe. The same goes for printing long strings.

(1) Open the PDF in Python using `pdfminer`. Put the entire text into a string called `pdftxt`.

<!-- grade_id: cell-3933a084b5a9d8c4, points: 0, grade: false, solution: true -->
```python
# your code here
from pdfminer.high_level import extract_text
pdftext = ''


```

<!-- grade_id: cell-903bdced1be3b890, points: 1, grade: true, solution: false -->
```python
# autograder tests - do not edit
print('Text is {} characters long'.format(len(pdftext)))
assert len(pdftext) == 526975
```

(2) If you take a look at the text, it seems that the different comments are separated by the `"Content"` string. 

Split your text into a list of comments using this separator. (This method isn't perfect, but it's a good first approximation.)

<!-- grade_id: cell-fdf68aab817288e8, points: 0, grade: false, solution: true -->
```python
comment_list = []  # your list of comments


```

<!-- grade_id: cell-50c524155e81eeb6, points: 1, grade: true, solution: false -->
```python
# autograder tests - do not edit

print(len(comment_list))
assert len(comment_list) == 193
```

(3) Now, let's clean up each of these comments. 

First, create a function that removes excess whitespace, punctuation, and stop words from a string. Also remove the following words that are in pretty much every comment: city, plan, Minneapolis, emailed, comments

Your function should take a string and return a cleaned-up string in the form of a list of words.

*Hint*: You'll first want to use regex to remove the excess whitespace and punctuation (and then whitespace again). Then create a list of words using `split()` and remove the stopwords. 

(Don't apply the function to your comment list yet — that's the next question.)

<!-- grade_id: cell-be7424cd658503f9, points: 0, grade: false, solution: true -->
```python
def clean_string(text):
    # your code here
    # it turns text (a string) into a cleaned list of words
    return cleaned_list_of_words



```

<!-- grade_id: cell-9441894209ca2a21, points: 1, grade: true, solution: false -->
```python
# autograder tests - do not edit
newstr = clean_string('A    very dirty 934\t999 string for minneapolis  plan IS like  this')

print(newstr)
assert newstr == ['dirty', 'string', 'like']
```

(4) Now, use your function to clean up the list of comments. Return a new list of cleaned comments.

*Hint*: a list comprehension is the simplest way to do this. But you could also do a regular `for` loop over your `comment_list`.

<!-- grade_id: cell-ec35dcae4c0ac7e0, points: 0, grade: false, solution: true -->
```python
cleaned_comment_list = []  # your code here


```

<!-- grade_id: cell-1b3b64b66a3db4ae, points: 1, grade: true, solution: false -->
```python
# autograder tests - do not edit
print(cleaned_comment_list[9][:10])

assert cleaned_comment_list[9][:10] == ['phase', 'super', 'americas', 'problems', 
                            'recent', 'threatened', 'revolt', 'license', 'issues', 'persist']
```

(5) Estimate an LDA topic model on your cleaned comment list. You could use the same parameters for `num_topics`, `alpha` and `eta` as we did in the video lectures. The challenge problem asks you to go deeper and experiment with different values, but feel free to do so here if you are inclined.

Visualize your topic model using `pyLDAvis`.

<!-- grade_id: cell-1a541320d85c3992, points: 2, grade: true, solution: true -->
```python
# your code here


```

(6) How do you interpret your topic model results? Explain in a sentence or two, or a few bullet points. 

<!-- grade_id: cell-da9d4db98f4b2f99, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

(7) Now, it's time for some sentiment analysis!

Write a function that calculates the sentiment score (polarity) for each comment in your (cleaned) list. 

The function should take a list of words and return a score.

*Hint*: Remember the `join()` function!

<!-- grade_id: cell-d7eb94c4d704bbeb, points: 0, grade: false, solution: true -->
```python
def get_sentiment(list_of_words):
    # your code here
    return polarity


```

<!-- grade_id: cell-0cdb129789f54d89, points: 1, grade: true, solution: false -->
```python
# autograder tests - do not edit
print(get_sentiment(['I', 'hate', 'the', 'idea', 'of', 'higher', 'densities']))
assert get_sentiment(['I', 'hate', 'the', 'idea', 'of', 'higher', 'densities'])==-0.275
```

(8) Now, apply the sentiment score to every string in your list (`cleaned_comment_list`). Create a new list of polarities. 

*Hint*: another list comprehension might be in order.

<!-- grade_id: cell-c19976e020fc0d16, points: 0, grade: false, solution: true -->
```python
sentiment_scores = []  # your list here


```

<!-- grade_id: cell-73ef071c85534bd3, points: 1, grade: true, solution: false -->
```python
# autograder tests - do not edit
import numpy as np
print(np.round(sentiment_scores[9], 2))
assert np.round(sentiment_scores[9], 2)==0.11
```

(8) Finally, plot a histogram of your scores. Make sure to add axis labels where appropriate.

<!-- grade_id: cell-e393336873100a86, points: 1, grade: true, solution: true -->
```python
# your code here


```

# Challenge Problem
Remember, you need to do at least two of these challenge problems this quarter.

This challenge problem is open ended for you to take in a direction that you are most interested in. Here are some suggestions (do 2-3 of these, or something analagous of your choice).

* For the topic modeling, experiment with `num_topics`, `alpha` and `eta` to get a meaningful set of topics. You might want to clean the data further as well, e.g. through lemmatizing and dropping other words (like we did for "Minneapolis")
* Analyze and plot the sentiment scores for posts that mention different keywords. Do you see a difference for those that mention "density," "zoning", "housing," "parking," etc.? Qualitatively, do the sentiment scores make intuitive sense (e.g. if you examine the most-positive and most-negative comments)?
* Until now, we've only considered the emailed comments. Compare the topics or sentiment scores to public comments received through a different channel ([all comments are here](https://minneapolis2040.com/received-public-comments/)). *Hint*: you could create a mega-function that takes the filename as the argument, and returns a set of topics and/or sentiment scores.
* Other ideas?

In all cases, write some brief interpretation in a markdown cell.

<!-- grade_id: cell-8f662277a6536325, points: 0, grade: true, solution: true -->
```python

```